# Assignment 4
## Description
In this assignment you must read in a file of metropolitan regions and associated sports teams from [assets/wikipedia_data.html](assets/wikipedia_data.html) and answer some questions about each metropolitan region. Each of these regions may have one or more teams from the "Big 4": NFL (football, in [assets/nfl.csv](assets/nfl.csv)), MLB (baseball, in [assets/mlb.csv](assets/mlb.csv)), NBA (basketball, in [assets/nba.csv](assets/nba.csv) or NHL (hockey, in [assets/nhl.csv](assets/nhl.csv)). Please keep in mind that all questions are from the perspective of the metropolitan region, and that this file is the "source of authority" for the location of a given sports team. Thus teams which are commonly known by a different area (e.g. "Oakland Raiders") need to be mapped into the metropolitan region given (e.g. San Francisco Bay Area). This will require some human data understanding outside of the data you've been given (e.g. you will have to hand-code some names, and might need to google to find out where teams are)!

For each sport I would like you to answer the question: **what is the win/loss ratio's correlation with the population of the city it is in?** Win/Loss ratio refers to the number of wins over the number of wins plus the number of losses. Remember that to calculate the correlation with [`pearsonr`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.pearsonr.html), so you are going to send in two ordered lists of values, the populations from the wikipedia_data.html file and the win/loss ratio for a given sport in the same order. Average the win/loss ratios for those cities which have multiple teams of a single sport. Each sport is worth an equal amount in this assignment (20%\*4=80%) of the grade for this assignment. You should only use data **from year 2018** for your analysis -- this is important!

## Notes

1. Do not include data about the MLS or CFL in any of the work you are doing, we're only interested in the Big 4 in this assignment.
2. I highly suggest that you first tackle the four correlation questions in order, as they are all similar and worth the majority of grades for this assignment. This is by design!
3. It's fair game to talk with peers about high level strategy as well as the relationship between metropolitan areas and sports teams. However, do not post code solving aspects of the assignment (including such as dictionaries mapping areas to teams, or regexes which will clean up names).
4. There may be more teams than the assert statements test, remember to collapse multiple teams in one city into a single value!

## Question 1
For this question, calculate the win/loss ratio's correlation with the population of the city it is in for the **NHL** using **2018** data.

```python
%%capture
#RUN FIRST, installs a missing library
import sys
!{sys.executable} -m pip install lxml==4.4.1
```

```python
import lxml
import sys
```

<!-- grade_id: cell-69b16e4386e58030, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
import scipy.stats as stats
import re

nhl_df=pd.read_csv("assets/nhl.csv")
cities=pd.read_html("assets/wikipedia_data.html")[1]
cities=cities.iloc[:-1,[0,3,5,6,7,8]]

def nhl_correlation(): 
    # YOUR CODE HERE
    raise NotImplementedError()
    
    population_by_region = [] # pass in metropolitan area population from cities
    win_loss_by_region = [] # pass in win/loss ratio from nhl_df in the same order as cities["Metropolitan area"]

    assert len(population_by_region) == len(win_loss_by_region), "Q1: Your lists must be the same length"
    assert len(population_by_region) == 28, "Q1: There should be 28 teams being analysed for NHL"
    
    return stats.pearsonr(population_by_region, win_loss_by_region)
```

<!-- grade_id: cell-ebe0b2dfe1067e63, points: 20, grade: true, solution: false -->
```python
nhl_df.head()
```

Output:
```
team                 GP                  W  \
0     Atlantic Division  Atlantic Division  Atlantic Division   
1  Tampa Bay Lightning*                 82                 54   
2        Boston Bruins*                 82                 50   
3  Toronto Maple Leafs*                 82                 49   
4      Florida Panthers                 82                 44   

                   L                 OL                PTS               PTS%  \
0  Atlantic Division  Atlantic Division  Atlantic Division  Atlantic Division   
1                 23                  5                113               .689   
2                 20                 12                112               .683   
3                 26                  7                105               .640   
4                 30                  8                 96               .585   

                  GF                 GA                SRS                SOS  \
0  Atlantic Division  Atlantic Division  Atlantic Division  
```

```python
cities.head(25)
```

Output:
```
Metropolitan area Population (2016 est.)[8]                   NFL  \
0            New York City                  20153634    GiantsJets[note 1]   
1              Los Angeles                  13310447  RamsChargers[note 4]   
2   San Francisco Bay Area                   6657982  49ersRaiders[note 6]   
3                  Chicago                   9512999         Bears[note 8]   
4        Dallas–Fort Worth                   7233323               Cowboys   
5         Washington, D.C.                   6131977              Redskins   
6             Philadelphia                   6070500                Eagles   
7                   Boston                   4794447     Patriots[note 14]   
8   Minneapolis–Saint Paul                   3551036               Vikings   
9                   Denver                   2853077               Broncos   
10   Miami–Fort Lauderdale                   6066387              Dolphins   
11                 Phoenix                   4661537             Cardinal
```

```python
nhl_df = nhl_df.query('year == 2018')
nhl_df = nhl_df[nhl_df['W'].str.isnumeric() & nhl_df['L'].str.isnumeric()].astype({'W':int, 'L':int}).reset_index(drop=True)
nhl_df
```

Output:
```
team  GP   W   L  OL  PTS  PTS%   GF   GA    SRS    SOS  \
0     Tampa Bay Lightning*  82  54  23   5  113  .689  296  236   0.66  -0.07   
1           Boston Bruins*  82  50  20  12  112  .683  270  214   0.62  -0.07   
2     Toronto Maple Leafs*  82  49  26   7  105  .640  277  232   0.49  -0.06   
3         Florida Panthers  82  44  30   8   96  .585  248  246  -0.01  -0.04   
4        Detroit Red Wings  82  30  39  13   73  .445  217  255  -0.48  -0.01   
5       Montreal Canadiens  82  29  40  13   71  .433  209  264  -0.68   0.00   
6          Ottawa Senators  82  28  43  11   67  .409  221  291  -0.85   0.00   
7           Buffalo Sabres  82  25  45  12   62  .378  199  280  -0.98   0.01   
8     Washington Capitals*  82  49  26   7  105  .640  259  239   0.21  -0.04   
9     Pittsburgh Penguins*  82  47  29   6  100  .610  272  250   0.23  -0.04   
10    Philadelphia Flyers*  82  42  26  14   98  .598  251  243   0.07  -0.03   
11  Columbus Blue Jackets*  82  45  30   7   97  .
```

```python
#retirando o asterisco
nhl_df['team'] = nhl_df['team'].str.replace(r'\*+','')

#separando os nomes dos times e criando uma nova coluna df com nome dos times
nhl_df['team_sep'] = nhl_df['team'].str.split(' ').str[-1]

#separando os nomes das cidades
#nhl_df['team_city'] = nhl_df['team'].str.split(' ')
nhl_df['team_city'] = nhl_df['team'].str.split(' ').apply(lambda x: x[0] if len(x) == 2 else ' '.join(x[0:2]))

```

Output:
```
C:\Users\jeanm\AppData\Local\Temp\ipykernel_18748\3172290646.py:2: FutureWarning: The default value of regex will change from True to False in a future version.
  nhl_df['team'] = nhl_df['team'].str.replace(r'\*+','')
```

```python
nhl_df
```

Output:
```
team  GP   W   L  OL  PTS  PTS%   GF   GA    SRS    SOS  \
0     Tampa Bay Lightning  82  54  23   5  113  .689  296  236   0.66  -0.07   
1           Boston Bruins  82  50  20  12  112  .683  270  214   0.62  -0.07   
2     Toronto Maple Leafs  82  49  26   7  105  .640  277  232   0.49  -0.06   
3        Florida Panthers  82  44  30   8   96  .585  248  246  -0.01  -0.04   
4       Detroit Red Wings  82  30  39  13   73  .445  217  255  -0.48  -0.01   
5      Montreal Canadiens  82  29  40  13   71  .433  209  264  -0.68   0.00   
6         Ottawa Senators  82  28  43  11   67  .409  221  291  -0.85   0.00   
7          Buffalo Sabres  82  25  45  12   62  .378  199  280  -0.98   0.01   
8     Washington Capitals  82  49  26   7  105  .640  259  239   0.21  -0.04   
9     Pittsburgh Penguins  82  47  29   6  100  .610  272  250   0.23  -0.04   
10    Philadelphia Flyers  82  42  26  14   98  .598  251  243   0.07  -0.03   
11  Columbus Blue Jackets  82  45  30   7   97  .591  242  23
```

```python
nhl_df['W/L'] = nhl_df['W'] / (nhl_df['W'] + nhl_df['L'])
nhl_df
```

Output:
```
team  GP   W   L  OL  PTS  PTS%   GF   GA    SRS    SOS  \
0     Tampa Bay Lightning  82  54  23   5  113  .689  296  236   0.66  -0.07   
1           Boston Bruins  82  50  20  12  112  .683  270  214   0.62  -0.07   
2     Toronto Maple Leafs  82  49  26   7  105  .640  277  232   0.49  -0.06   
3        Florida Panthers  82  44  30   8   96  .585  248  246  -0.01  -0.04   
4       Detroit Red Wings  82  30  39  13   73  .445  217  255  -0.48  -0.01   
5      Montreal Canadiens  82  29  40  13   71  .433  209  264  -0.68   0.00   
6         Ottawa Senators  82  28  43  11   67  .409  221  291  -0.85   0.00   
7          Buffalo Sabres  82  25  45  12   62  .378  199  280  -0.98   0.01   
8     Washington Capitals  82  49  26   7  105  .640  259  239   0.21  -0.04   
9     Pittsburgh Penguins  82  47  29   6  100  .610  272  250   0.23  -0.04   
10    Philadelphia Flyers  82  42  26  14   98  .598  251  243   0.07  -0.03   
11  Columbus Blue Jackets  82  45  30   7   97  .591  242  23
```

```python
nhl_df
```

## Question 2
For this question, calculate the win/loss ratio's correlation with the population of the city it is in for the **NBA** using **2018** data.

<!-- grade_id: cell-5a5f21279e3d3572, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
import scipy.stats as stats
import re

nba_df=pd.read_csv("assets/nba.csv")
cities=pd.read_html("assets/wikipedia_data.html")[1]
cities=cities.iloc[:-1,[0,3,5,6,7,8]]

def nba_correlation():
    # YOUR CODE HERE
    raise NotImplementedError()
    
    population_by_region = [] # pass in metropolitan area population from cities
    win_loss_by_region = [] # pass in win/loss ratio from nba_df in the same order as cities["Metropolitan area"]

    assert len(population_by_region) == len(win_loss_by_region), "Q2: Your lists must be the same length"
    assert len(population_by_region) == 28, "Q2: There should be 28 teams being analysed for NBA"

    return stats.pearsonr(population_by_region, win_loss_by_region)
```

<!-- grade_id: cell-e573b2b4a282b470, points: 20, grade: true, solution: false -->
```python

```

## Question 3
For this question, calculate the win/loss ratio's correlation with the population of the city it is in for the **MLB** using **2018** data.

<!-- grade_id: cell-33b00fc3f3467b0c, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
import scipy.stats as stats
import re

mlb_df=pd.read_csv("assets/mlb.csv")
cities=pd.read_html("assets/wikipedia_data.html")[1]
cities=cities.iloc[:-1,[0,3,5,6,7,8]]

def mlb_correlation(): 
    # YOUR CODE HERE
    raise NotImplementedError()
    
    population_by_region = [] # pass in metropolitan area population from cities
    win_loss_by_region = [] # pass in win/loss ratio from mlb_df in the same order as cities["Metropolitan area"]

    assert len(population_by_region) == len(win_loss_by_region), "Q3: Your lists must be the same length"
    assert len(population_by_region) == 26, "Q3: There should be 26 teams being analysed for MLB"

    return stats.pearsonr(population_by_region, win_loss_by_region)
```

<!-- grade_id: cell-764d4476f425c5a2, points: 20, grade: true, solution: false -->
```python
import pandas as pd
import numpy as np
import scipy.stats as stats
import re

mlb_df=pd.read_csv("assets/mlb.csv")
cities=pd.read_html("assets/wikipedia_data.html")[1]
cities=cities.iloc[:-1,[0,3,5,6,7,8]]

def mlb_correlation(): 
    mlb_df = pd.read_csv("assets/mlb.csv")
    cities = pd.read_html("assets/wikipedia_data.html")[1]
    cities = cities.iloc[:-1, [0, 3, 5, 6, 7, 8]]

    # mlb
    mlb_df = mlb_df.query('year == 2018')
    mlb_df = mlb_df.replace(
        {'Boston Red Sox': 'RedSox', 'Chicago White Sox': 'WhiteSox'})
    mlb_df['team'] = mlb_df['team'].str.split(' ').str[-1]

    # cities
    teams = {'YankeesMets': 'Yankees Mets', 'DodgersAngels': 'Dodgers Angels',
            'GiantsAthletics': 'Giants Athletics', 'CubsWhite Sox': 'Cubs WhiteSox', 'Red Sox': 'RedSox'}
    cities['MLB'] = cities['MLB'].str.replace(r'\[.*\]', '', regex=True)
    cities['MLB'] = cities['MLB'].replace(teams).str.split(' ')
    cities = cities.explode('MLB')

    # merge
    df = pd.merge(mlb_df, cities, left_on='team', right_on='MLB')
    df = df[['Metropolitan area', 'Population (2016 est.)[8]', 'MLB', 'W-L%']].astype(
        {'Population (2016 est.)[8]': 'int64', 'W-L%': 'float'})

    # mean and merge
    df_mean = df.groupby('Metropolitan area')[['W-L%']].mean()
    df = pd.merge(df, df_mean, on='Metropolitan area').drop_duplicates(
        'Metropolitan area').reset_index(drop=True)
    df = df[['Metropolitan area',
            'Population (2016 est.)[8]', 'MLB', 'W-L%_y']].rename(columns={'W-L%_y': 'W-L%'})

    # pass in metropolitan area population from cities
    population_by_region = df['Population (2016 est.)[8]']
    # pass in win/loss ratio from mlb_df in the same order as cities["Metropolitan area"]
    win_loss_by_region = df['W-L%']
    assert len(population_by_region) == len(win_loss_by_region), "Q3: Your lists must be the same length"
    assert len(population_by_region) == 26, "Q3: There should be 26 teams being analysed for MLB"
    corr = stats.pearsonr(population_by_region, win_loss_by_region)

    return corr[0]
    raise NotImplementedError()
```

```python
mlb_correlation()
```

Output:
```
0.15003737475409495
```

## Question 4
For this question, calculate the win/loss ratio's correlation with the population of the city it is in for the **NFL** using **2018** data.

<!-- grade_id: cell-8ccebc209aeec8d9, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
import scipy.stats as stats
import re

nfl_df=pd.read_csv("assets/nfl.csv")
cities=pd.read_html("assets/wikipedia_data.html")[1]
cities=cities.iloc[:-1,[0,3,5,6,7,8]]

def nfl_correlation(): 
    # YOUR CODE HERE
    raise NotImplementedError()
    
    population_by_region = [] # pass in metropolitan area population from cities
    win_loss_by_region = [] # pass in win/loss ratio from nfl_df in the same order as cities["Metropolitan area"]

    assert len(population_by_region) == len(win_loss_by_region), "Q4: Your lists must be the same length"
    assert len(population_by_region) == 29, "Q4: There should be 29 teams being analysed for NFL"

    return stats.pearsonr(population_by_region, win_loss_by_region)
```

<!-- grade_id: cell-de7b148b9554dbda, points: 20, grade: true, solution: false -->
```python

```

## Question 5
In this question I would like you to explore the hypothesis that **given that an area has two sports teams in different sports, those teams will perform the same within their respective sports**. How I would like to see this explored is with a series of paired t-tests (so use [`ttest_rel`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.ttest_rel.html)) between all pairs of sports. Are there any sports where we can reject the null hypothesis? Again, average values where a sport has multiple teams in one region. Remember, you will only be including, for each sport, cities which have teams engaged in that sport, drop others as appropriate. This question is worth 20% of the grade for this assignment.

<!-- grade_id: cell-92f25f44b8d1179f, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
import scipy.stats as stats
import re

mlb_df=pd.read_csv("assets/mlb.csv")
nhl_df=pd.read_csv("assets/nhl.csv")
nba_df=pd.read_csv("assets/nba.csv")
nfl_df=pd.read_csv("assets/nfl.csv")
cities=pd.read_html("assets/wikipedia_data.html")[1]
cities=cities.iloc[:-1,[0,3,5,6,7,8]]

def sports_team_performance():
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Note: p_values is a full dataframe, so df.loc["NFL","NBA"] should be the same as df.loc["NBA","NFL"] and
    # df.loc["NFL","NFL"] should return np.nan
    sports = ['NFL', 'NBA', 'NHL', 'MLB']
    p_values = pd.DataFrame({k:np.nan for k in sports}, index=sports)
    
    assert abs(p_values.loc["NBA", "NHL"] - 0.02) <= 1e-2, "The NBA-NHL p-value should be around 0.02"
    assert abs(p_values.loc["MLB", "NFL"] - 0.80) <= 1e-2, "The MLB-NFL p-value should be around 0.80"
    return p_values
```

<!-- grade_id: cell-fb4b9cb5ff4570a6, points: 20, grade: true, solution: false -->
```python

```